---
title: "Working with R"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
gradethis_setup()
knitr::opts_chunk$set(echo = FALSE)
```

## Variables

Let's see how variables are used.

Write the R code required to add two plus two:

```{r two, exercise=TRUE}

```
```{r two-check}
grade_result(
  pass_if(~ identical(.result, 4), "Good job!")
)
```

Now create two variables `a` and `b`, each containing the number 2 and add those

```{r two-plus-two-variables, exercise=TRUE}

```
```{r two-plus-two-variables-solution}
a <- 2
b <- 2
a + b
```
```{r two-plus-two-variables-check}
grade_code()
```

## Vectors

Variables can contain more than just simple numbers, try assigning a vector of 16 random uniform numbers from the `runif()` function to a variable called `rand`.

```{r rand, exercise=TRUE}

```
```{r rand-solution}
rand <- runif(16)
```
```{r rand-check}
grade_code()
```

And vectors can be operated on as a single thing, sometimes the operation will return a single number, sometimes a vector.

Here are some R vectors.
```
a <- c(1,2,3)
b <- c("a", "b")
d <- c(a,b)
```

```{r vector-ops-1}
question(
  "What is the result of `mean(a)`",
  answer("6"), answer("2", correct = TRUE), answer("3")
)
```
```{r vector-ops-2}
question(
  "What is the result of `a + 2`?",
  answer("3,4,5", correct = TRUE),
  answer("1,2,3,2")

)
```
```{r vector-ops-3}
question(
  "What is the result of `d * 2`?",
  answer("3,4,5"),
  answer("2,4,6"),
  answer("Error ...", correct = TRUE),
  answer("2,4,6,NA,NA")
)

```

## Dataframes

Dataframes are rectangular spreadsheet style objects that hold data we wish to analyse. There is a small one called `small_df` loaded into memory for you, the function `str()` tells us about a dataframe. Use `str()` on the dataframe `small_df`

```{r, str-dataframe-setup, echo=FALSE}
small_df <- data.frame(
    names = c("Guido", "Marty", "Alan"),
    age = c(24,45,11),
    score = runif(3) * 100
  )

```
```{r, str-dataframe, exercise = TRUE}

```
```{r str-dataframe-solution}
str(small_df)
```
```{r str-dataframe-check}
grade_code()
```

```{r, dataframe}
question(
  "What does the output tell us about the column `names`?",
  answer("it has chromosome data in it"),
  answer("it is text-based (CHaRacter) data", correct = TRUE),
  answer("it has numeric data in it")
)
```

## Using R Help

Use the help function to answer these questions

```{r, help-1}
question(
  "Does the `mean()` function compute",
  answer("the geometric mean"),
  answer("the average"),
  answer("the arithmetic mean"),
  answer("the harmonic mean"),
  answer("the (trimmed) arithmetic mean", correct = TRUE)
)
```

```{r, help-2}
question(
  "What does the `c` of the `c()` function stand for?",
  answer("concatanate"),
  answer("convolve"),
  answer("combine", correct = TRUE),
  answer("complete")
)
```

```{r, help-3}
question(
  "Do you fully understand what the `apply` function does?",
  answer("Yes!"),
  answer("No"),
  answer("No, but that's ok, I've read the help and it's a bit opaque, but I think some googling will help", correct = TRUE)
  )

```


